<img style="float: left;" src="earth-lab-logo-rgb.png" width="150" height="150" />

# Earth Analytics Education - Bootcamp Course Fall 2020

## Important  - Assignment Guidelines

1. Before you turn in your assignment, make sure to run the entire notebook with a fresh kernel. To do this first, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart & Run All)
2. In the cells below you will replace the `raise NotImplementedError()` code with your code that addresses the activity challenge. If you don't replace that code, your notebook will not run properly.

```
# YOUR CODE HERE
raise NotImplementedError()
```

3. Any open ended questions will have a "YOUR ANSWER HERE" within a markdown cell. Replace that text with your answer also formatted using Markdown.
4. **IMPORTANT: DO NOT RENAME THIS NOTEBOOK!** If the file name changes, the autograder will not grade your assignment properly.
5. **Do not rename the notebook file.** If you do, the autograder will not recognize your submisson.
6. When you plot, please comment out `plt.show()` as the code below will effectively run `plt.show()` for you and also will grab your plot for autograding. DO NOT DELETE any code that says `DO NOT REMOVE LINE BELOW`. That code is for autograding!!

```
### DO NOT REMOVE LINE BELOW ###
student_plot1_ax = nb.convert_axes(plt)
```



## Follow to PEP 8 Syntax Guidelines

* Run the `autopep8` tool on all cells prior to submitting (HINT: hit shift + the tool to run it on all cells at once!
* Use clear and expressive names for variables. 
* Organize your code to support readability.
* Check for code line length
* Use comments and white space sparingly where it is needed


### Add Your Name Below 
**Your Name:**

<img style="float: left;" src="colored-bar.png"/>

---

# Week 6 Homework Template - Numpy Arrays

To complete assignment 5, be sure you have reviewed Chapters 11 and 12 in Section 4 and Chapter 14 in Section 6 of the <a href="https://www.earthdatascience.org/courses/intro-to-earth-data-science/" target="_blank">Intro to Earth Data Science online textbook</a> online textbook, which cover working with packages and paths as well as the numpy package.

Read the instructions for each question carefully to successfully complete the required tasks.

## Assignment Data

For this assignment, you will write **Python** code to download and plot average monthly temperatures in 2017 for Boulder, CO <a href="https://www.esrl.noaa.gov/psd/boulder/Boulder.mm.html" target="_blank">U.S. National Oceanic and Atmospheric Administration (NOAA)</a>. Then you will process some NOAA precipitation data also collected in Boulder, Colorado.

```python
# Core imports needed for grading - Do not modify this cell!
import matplotcheck.notebook as nb
from matplotcheck.base import PlotTester
import pandas as pd
```

## Import Python Packages

In the cell below, add code **after the line for `Your Code Here`**, replacing `raise NotImplementedError()` with your code to import the package/module needed to:
* create plots (hint: you have only learned about one so far!)
* set your working directory
* download data using earthpy functions
* work with numpy arrays

Be sure to list the package imports following the appropriate PEP 8 order. 

<!-- grade_id: student-imports-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: student-imports-tests, points: 4, grade: true, solution: false -->
```python
# Test package imports - DO NOT MODIFY THIS CELL!
import_answer_points = 0

try:
    np.nan
    print("\u2705 Score! Numpy has been imported as a np!")
    import_answer_points += 1
except NameError:
    print("\u274C Numpy has not been imported as a np, please make sure to import is properly.")

try:
    plt.show()
    print("\u2705 Nice! matplotlib.pyplot has been imported as plt!")
    import_answer_points += 1
except NameError:
    print("matplotlib.pyplot has not been imported as plt, please make sure to import is properly.")

try:
    os.getcwd()
    print("\u2705 Great work! The os module has imported correctly!")
    import_answer_points += 1
except NameError:
    print("\u274C Oops make sure that the os package is imported.")

try:
    data = et.io
    print("\u2705 Great work! The earthpy package has imported correctly!")
    import_answer_points += 1
except NameError:
    print("\u274C Oops make sure that the earthpy package is imported using the alias et.")

print("\n \u27A1 You received {} out of 4 points.".format(import_answer_points))

import_answer_points
```

## Dataset 1: Download Temperature Data Using EarthPy

In the cell below, add code **after the line for `Your Code Here`**, replacing `raise NotImplementedError()` with your code, to complete the following task:

* Download the following datasets: 
    * months.txt from https://ndownloader.figshare.com/files/12565619
        * this file contains abbreviated month names as text strings
    * boulder-temp-2017.txt from https://ndownloader.figshare.com/files/12767966
        * this file contains average monthly temperature (Fahrenheit) in 2017 for Boulder, Colorado

<!-- grade_id: student-calc-variables-answer, points: 8, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Set Your Working Directory

In the cell below, add code **after the line for `Your Code Here`**, replacing `raise NotImplementedError()` with your code, to set your working directory to **your `earth-analytics` directory in your home directory**.

Be sure to use the appropriate functions that will allow your code to run successfully on any operating system.

<!-- grade_id: student-create-list-mean-fire-size-answer, points: 8, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<img style="float: left;" src="colored-bar.png"/>

# Challenge - Create Paths to Files

Create a path variable called `boulder_temp_path` that directs python to the 
`boulder-temp-2017.txt` file using `os.path.join`.
The path should assume that you working directory is **earth-analytics** and 
should be RELATIVE to the earth-analytics directory
 
Create a second path variable that directs Python to the **month.txt** file. 
Name that variable: `months_path`.

<!-- grade_id: student-set-data-paths, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test-data-paths, points: 4, grade: true, solution: false -->
```python
# DO NOT MODIFY ANY CODE IN THIS CELL  - TESTS THAT THE DATA PATHS ARE SET CORRECTLY
path_set_points = 0
try:
    if not os.path.isfile(boulder_temp_path):
        print("Variable boulder_temp_path exists, but is not a valid file path. Make sure the path you entered points to the file, and is formatted like a path using os.path.join()")
    else:
        if os.path.basename(os.path.normpath(boulder_temp_path)) == 'boulder-temp-2017.txt':
            print(
                "\u2705 Variable boulder_temp_path exists and points to the right file!")
            path_set_points += 2
        else:
            print(
                "\u274C Variable boulder_temp_path exists, but points to the wrong file.")
except NameError:
    print("\u274C Variable boulder_temp_path wasn't assigned.")

try:
    if not os.path.isfile(months_path):
        print("\u274C Variable months_path exists, but is not a valid file path. Make sure the path you entered points to the file, and is formatted like a path using os.path.join()")
    else:
        if os.path.basename(os.path.normpath(months_path)) == 'months.txt':
            print(
                "\u2705 Variable months_path exists exists and points to the right file!")
            path_set_points += 2
        else:
            print("\u274C Variable months_path exists, but points to the wrong file.")
except NameError:
    print("\u274C Variable months_path wasn't assigned.")

print("\n \u27A1 You received {} out of 4 points.".format(path_set_points))

path_set_points
```

<img style="float: left;" src="colored-bar.png"/>

## Import Data Into Numpy Arrays

In the cell below, add your code **after the line for `Your Code Here`**, replacing `raise NotImplementedError()` with your code, to import the downloaded text data files into individual numpy arrays using **relative paths**. 

Your code should create two individual numpy arrays: one for month names and one for the temperature data. 

Be sure to use the appropriate functions that will allow your code to run successfully on any operating system.

<!-- grade_id: student-create-list-number-fires-answer, points: 12, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Challenge Figure 1: Create Figure From Numpy Arrays

In the cell below, replace `raise NotImplementedError()` with your code, to complete the following task:

* Create a bar plot that plots the two numpy arrays that you opened up above. 
    * The x axis should contains months, 
    * The y axis should contain temperature values.   
* Modify the bar plot color. 
* Include an appropriate title and axes labels. Include units of measurement where appropriate. 

For your title and labels, be sure to think about the following pieces of information that could help someone easily interpret the plot:
* geographic coverage or extent of data.
* duration or temporal extent of the data.
* what was actually measured and/or represented by the data.
* units of measurement.

**IMPORTANT:** Comment out the code line `plt.show()` in your plot code to avoid an additional blank subplot.

<!-- grade_id: student-create-plot-1-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

### DO NOT REMOVE LINE BELOW ###
student_plot1_ax = nb.convert_axes(plt)
```

<!-- grade_id: test-plot-1-answer, points: 15, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL - Test For Your Plot

# This will test that your plot above was created correctly
# Remember that plt.show() should be commented out for this to work
plot_1_check = PlotTester(student_plot1_ax)
bar_plot = False

try:
    plot_1_check.assert_plot_type('bar')
    print("\u2705 Successfully created a plot of bar type!")
    bar_plot = True
except AssertionError:
    print("\u274C The plot is not a bar plot, make sure to make it to a bar plot. Tests will not run successfully if this does not pass.")


```

## Discussion Figure 1: Explain Your Plot

In the Markdown cell below, answer the following questions about your plot using a **bullet list**.

1. What was the warmest month in 2017?
2. What was the coolest month in 2017?
3. The data provided are only for one year (2017). What other data would be useful to add to a plot like this to begin to evaluate whether 2017 was a warm year compared to previous years? 

Remove any existing text in the cell below before adding your answer.

<!-- grade_id: student-explain-pep-8, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Dataset 2: Download Monthly Precipitation Data

In the cell below complete the following task:

* Download the `monthly-precip-mm-1998-to-2017.csv` using the url: https://ndownloader.figshare.com/files/12799931 . This **.csv** file 
contains average monthly precipitation (millimeters) for the years 1998 to 
2017 for Boulder, Colorado.
* Then open the data as a numpy array using `np.loadtxt()`

<!-- grade_id: student-boulder-precip, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Check Shape of Numpy Array

In the cell below, add code **after the line for `Your Code Here`**, replacing `raise NotImplementedError()` with your code, to check the shape of the numpy array containing the average monthly precipitation data. 

<!-- grade_id: student-check-shape, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test-shape-correct, points: 2, grade: true, solution: false -->
```python
# Test that the tuple output of shape produced above is correct

```

## Run Summary Statistics on Numpy Array

In the cell below:

* Calculate the mean of the precipitation values for each year (i.e. across all months in each year).
* Calculate the median precipitation value for each year

Once you have these values, create a python list `[]` that contains the numpy array 
for mean values first (at the first index) and the median values numpy array second.

At the end of your cell, be sure to call the newly created list object so 
the output prints below (do not use the print statement!). Just call the variable
name at the end of the cell. 


<!-- grade_id: student-get-mean-median, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test-mean-median-1, points: 2, grade: true, solution: false -->
```python
student_list = _
list_points = 0

if isinstance(student_list, list):
    print("\u2705 You created a list, well done!")
    list_points += 1
else:
    print("\u274C The variable you made isn't a list, please assign it to be one.")

if student_list[0].dtype == 'float64':
    print("\u2705 The data stored in the numpy array inside your lists is in the expected format: float64!")
    list_points += 1
else:
    print("\u274C The data stored inside your numpy arrays in your list is not in the expected format: float64.")

list_points
```

<!-- grade_id: test-mean-median-2, points: 4, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL

```

## Run Calculations on Numpy Arrays

Above, you were working with precipitation data collected in millimeters. 

Create a new list that contains median and mean precipitation values in INCHES.

* Convert the values in the **numpy** arrays you created above (i.e. the median and mean of the precipitation values for each month) **from millimeters to inches** (1 inch = 25.4 millimeters). 
* Use the appropriate **assignment operator** to convert both arrays. Think about which mathematical operation (e.g. `+`, `*`) you need to run this calculation. 

Be sure to call the final list variable at the end of the cell. 


<!-- grade_id: student-convert-units, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# visible test that the output above is a list object

student_list = _

if isinstance(student_list, list):
    print("\u2705 Yes! You have created a list above - great work!")
else:
    print("\u274C Oops, please make sure that you've outputted a list above with "
          "with two numpy  arrays in it")
```

<!-- grade_id: test-list-of-arrays, points: 4, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL

```

<img style="float: left;" src="colored-bar.png"/>

## Challenge Figure 2: Create Multi-plot Figure From Arrays (20 points)

In the cell below complete the following tasks:

* Create one multi-plot figure that contains two subplots that are **side by side**:
    * top plot: median of the precipitation values for each month in inches as a scatter plot
    * bottom plot: mean of the precipitation values for each month in inches as a scatter plot
* **Use a different color for each plot and adjust the marker styles.
    * For each plot, be sure to include appropriate titles and axes labels including units of measurement where appropriate. 
* Add an overall title for the entire figure. 

For your titles and labels, be sure to think about the following pieces of information that could help someone easily interpret the plot:
* geographic coverage or extent of data.
* duration or temporal extent of the data.
* what was actually measured and/or represented by the data.
* units of measurement.

Remember to always comment out the code line `plt.show()` in your plot code. 
If you do not comment out that code line, you may see an extra empty plot 
underneath your desired figure. 

HINT: Are your labels overlapping? See what happens when you add `plt.tight_layout()` above the 
**### DO NOT REMOVE LINE BELOW ###** part of your plot 
code cell below

<!-- grade_id: figure-2-median-mean-plots-solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

### DO NOT REMOVE LINE BELOW ###
numpy_2_plot = nb.convert_axes(plt, which_axes="all")
```

<!-- grade_id: figure-2-hidden-tests, points: 20, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL
```

<img style="float: left;" src="colored-bar.png"/>

## Discuss Your Workflow In This Notebook

In the Markdown cell below, answer the following questions using a **numbered list**:

1. List two elements in this Jupyter Notebook file that supports reproducibility.

Remove any existing text in the cell before adding your answer.

<!-- grade_id: student-explain-reproduc-answer, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

<img style="float: left;" src="colored-bar.png"/>

# PEP 8, Syntax and Spelling / Grammar

* the notebook starts at [1] and runs all the way through (2 points)
* PEP 8 syntax is followed for code format, comments, imports, etc (3 points)
* Spelling, grammar and good use of Markdown syntax are considered for all text answers. (2 points)

## OPTIONAL Bonus Challenge Plot: Create Figure with Loops (10 points)

**This is an optional bonus plot if you want a challenge this week!**

In the cell below:

- Create a plot with lines showing the monthly precipitation for each year in the `monthly-precip-mm-1998-to-2017.csv` dataset. In your plot, be sure to do the following:
    - **Use a different color for each line representing one year in the dataset.**
    - Include appropriate titles and axes labels including units of measurement where appropriate.
    - Add a legend to the plot showing what color line represents each year.

### Using For Loops to Plot

In order to completely this task, you will need to write a for loop 
in your code to assign unique color and year values to each line in the plot. 

1. **Create a list of Years** First, create a list of years that the data covers. The file name tells you that 
the data ranges from 1998 to 2017. In order to assign a year to each row in the 
`numpy` array that you have our data in, you need to make a list of years that 
covers that range. 

2. **Create a list of Colors** Create a list of unique colors to assign to each line in the plot 
(each line represents data for one year). You can do this manually, 
(i.e., `['blue', 'green', ... , 'red']`) or you can do this using `seaborn`, 
a package that is used to style plots in **Python**. Using `seaborn`, you can 
create a list of unique color values from a colormap (`cmap`). For example, 
`sns.color_palette("mako", n_colors=10)` will return you a list of 10 unique colors 
found within the `mako` color map. How many unique colors it returns depends on 
what you set the `n_colors` argument to. You can use this to get the number of 
unique colors you need for your plot.

3. **Write a for loop** You will need to loop through all of these lists alongside your 
`boulder_precip` data in order to plot all of the years individually. You 
can loop through multiple lists at a time with the help of the `zip` function 
found in default **Python**. This will combine all of your lists into one `zip` 
object, which can be looped through. For example, if you have three lists, `list_a`, 
`list_b`, and `list_c`, and you wanted to loop through all three of them simultaneously, 
you could use `zip` to do this:

```
for var_a, var_b, var_c in zip(list_a, list_b, list_c):
    print(var_a, var_b, var_c)
```

Here, each variable `var_x` is the variable found in each respective list at that index. You can use this principle in combination with `ax.plot()` to plot each line in the `boulder_precip` dataset by year and with a unique color. Set the x data to a list of months, the y data to the data in the `boulder_precip` dataset, the `label` argument to the year, and the `color` argument to the unique color for each year. 

4. **Create a legend** Legends can be tricky in python. To create a legend you need to do two things
   * First, you need to assigne a label value - this will create the text for each line in your plot. In this case each line represents a year so label will equal a string with the year in it for each iteration of your for loop.
   * Finally you need to call `plt.legend()` at the end of your plot code to tell Python to render your legend. 
   
Your code will looks something like this

for ayear, acolor, anarray in zip(years_list, colors_list, numpyarray):
    ax.plot(months, # the x axis months
            anarray, # the array dimension being iterated on
            label=ayear, # the year being iterated on
            color=acolor) # The color that is being iterated on

*****
Remember to always comment out the code line `plt.show()` in your plot code. If you do not comment out that code line, you may see an extra empty plot underneath your desired figure and it will cause issues with the autograder.

HINT: Are your labels overlapping? See what happens when you add `plt.tight_layout()` above the ### DO NOT REMOVE LINE BELOW ### part of your plot code cell below

***

HINT - run the code below to see how "zip" works in Python

```python
colors = ["red", "green", "blue"]
years = ["2015", "2016", "2017"]

for acolor, ayear in zip(colors, years):
    print(acolor)
    print(ayear)
```
***

<!-- grade_id: cell-bonus-plot-solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

### DO NOT REMOVE LINE BELOW ###
numpy_3_plot = nb.convert_axes(plt)
```

<!-- grade_id: cell-bonus-figure-tests, points: 0, grade: true, solution: false -->
```python

```